<a href="https://colab.research.google.com/github/raydot/coursera/blob/main/UCB_Week_Two_Minheaps.ipynb" target="_parent"><img src="https://colab.research.google.com/assets/colab-badge.svg" alt="Open In Colab"/></a>

# Problem Set \# 2 (Basic Data Structures and Heaps)

Topics covered:
  - Basic data-structures
  - Heap data-structures
  - Using heaps and arrays to realize interesting functionality.

 #### Pseudocode for recursive BubbleUp:

(j // 2) is Pythonic for ⌊j / 2⌋

`left(i): 2i`
`right(i): 2i + 1`
``` python
bubbleUp(A, j)
  if j ≤ 1:
    return
  else
  if (A[j] < A[j//2):
    swap(A[j], A(j//2])
    bubbleUp(A(j//2))
```

 #### Pseudocode for recursive BubbleDown:

```python
bubbleDown(A, j):
  # n = length of A
  if (left(i) > n) # A has no children
    return
  if (left(i) ≤ n and right(i) > n): # A has one child
    if (A[i] > A[left(i)]):
      swap(A[i], A[left(i)])
      bubbleDown(A, left(i)
  else # A has two children
    #locate smaller child
    if (A([left(i)] < A[right(i)]:
      small := left(i)
    else
      small := right(i)
    swap (A[i], A[small])
    bubbleDown(A, small)
  
```

## Problem 1 (Least-k Elements Datastructure)


We saw how min-heaps can efficiently allow us to query the least element in a heap (array). We would like to modify minheaps in this exercise to design a data structure to maintain the __least k__ elements for a  given $k \geq 1$ with $$k = 1$$ being the minheap data-structure.

Our design is to hold two arrays:
  - (a) a sorted array `A` of $k$ elements that forms our least k elements; and  
  - (b) a minheap `H` with the remaining $n-k$ elements.  
**A = least k elements**<br />
**H = minheap of remaining k elements**<br />
**n = total number of elements**<br />
**k = # of elements in A**

Our data structure will itself be a pair of arrays `(A,H)` with the following property:
 - `H` must be a minheap
 - `A` must be sorted of size $k$.
 - Every element of `A` must be smaller than every element of `H`.

The key operations to implement in this assignment include:
  - insert a new element into the data-structure
  - delete an existing element from the data-structure.


We will first ask you to design the data structure and them implement it.

### (A) Design Insertion  Algorithm

Suppose we wish to insert a new element with key $j$ into this data structure. Describe the pseudocode. Your pseudocode must deal with two cases: when the inserted element $j$ would be one of the `least k` elements i.e, it belongs to the array `A`; or when the inserted element belongs to the heap `H`.

How would you distinguish between the two cases?

_Belongs in A if it's smaller than the biggest number in A.  Otherwise, belongs in H._

- You can assume that heap operations such as `insert(H, key)` and `delete(H, index)` are defined.
- Assume that the heap is indexed as  `H[1]`,...,`H[n -k]` with `H[0]` being unused. <br />*_Why is `H[0]` unused?_*  **To support sending an element out of H.  THIS PROPERTY MUST BE MAINTAINED!**
- Assume $ n > k$, i.e, there are already more than $k$ elements in the data structure.

- What is the complexity of the insertion operation in the worst case in terms of $k, n$.  __DO NOT SKIP THIS!__

__Unfortunately, we cannot grade your answer. We hope you will use this to design your datastructure on paper before attempting to code it__

#### Heap Operations (`insert, delete, heapify, heapsort`)
  
**Insert:**

- Append to end `A.append(e)`
- BubbleUp

Runtime complexity is $\Theta$(log<sub>n</sub>)

**Delete:**
  - If item is at end, remove and readjust array size to `n - 1`
  - if in middle:
    1.   replace A<sub>j</sub> with A<sub>n</sub>
    2.   adjust size to n - 1
    3.   fix what's broken:
      - BubbleUp if (A<sub>n</sub>) <sup>in jth pos</sup> is < parent
      - BubbleDown if A<sub>n</sub> > one or both children

Runtime complexity is $\Theta$(log<sub>n</sub>)

**Heapify:**

BubbleDown each element, moving from end to start.  Can start at n/2 because anything "below" that is a leaf.  So it's simply:

```python
  for i = n/2 down to 1
    bubbleDown(A, i)
```
Runtime complexity is $\Theta$(n)

**Heapsort:**

(Because a heap is not necessarily a sorted array)

  1. Heapify(A)
  2. ```python
  for i = 1 to n
        insert A[1] into result
        delete(A[1])
      ```

This can be done into a new array `[result]`, or in place by appending the removed element to the existing array.  

Runtime complexity is $\Theta$(n log n)

<!-- grade_id: cell-c21e18266dbaa2b6, points: 0, grade: true, solution: true -->


```
Insert a new element with key j into the data structure
if j is a least k element {
  insert j
  put removed element into H
} else if it isnt {
  insert into heap
  heapify heap
}
```

YOUR ANSWER HERE


### (B) Design Deletion Algorithm

Suppose we wish to delete an index $j$ from the top-k array $A$. Design an algorithm to perform this deletion. Assume that the heap is not empty, in which case you can assume that the deletion fails.



- You can assume that heap operations such as `insert(H, key)` and `delete(H, index)` are defined.
- Assume that the heap is indexed as  `H[1]`,...,`H[n -k]` with `H[0]` being unused (for item deletion).
- Assume $ n > k$, i.e, there are already more than $k$ elements in the data structure.

What is the complexity of the insertion operation in the worst case in terms of $k, n$.

__Unfortunately, we cannot grade your answer. We hope you will use this to design your data structure on paper before attempting to code it up__

<!-- grade_id: cell-691fd6725fc83928, points: 0, grade: true, solution: true -->
YOUR ANSWER HERE
For each item in the array it has to:

1.   Move the last item to the front
2.   Pop off the last item
3.   Bubble down

### Answer from down-a-below

* First, in order to delete the index j from array, move elements from j+1 .. k-1 left one position.
* Insert the minimum heap element at position  𝑘−1  of the array A.
* Delete the element at index 1 of the heap.

## (C) Program your solution by completing the code below

Note that although your algorithm design above assume that your are inserting and deleting from cases where $n \geq k$, the data structure implementation below must handle $n < k$ as well. We have provided implementations for that portion to help you out.

<!-- grade_id: cell-4359c4044763e75a, points: 0, grade: false, solution: true -->
```python
# First let us complete a minheap data structure.
# Please complete missing parts below.

class MinHeap:
    def __init__(self):
        # H = the heap of the N-K elements.
        self.H = [None]

    def size(self):
        return len(self.H)-1

    def __repr__(self):
        return str(self.H[0:])

    def satisfies_assertions(self):
        for i in range(2, len(self.H)):
            assert self.H[i] >= self.H[i//2],  f'Min heap property fails at position {i//2}, parent elt: {self.H[i//2]}, child elt: {self.H[i]}'

    def is_empty(self):
        print(f"IS EMPTY? {self.H}")
        return self.H == [None]

    def min_element(self):
        return self.H[1]

    ## bubble_up function at index
    ## WARNING: this function has been cut and paste for the next problem as well
    def bubble_up(self, index):
        # print("HELLO FROM BUBBLE_UP at index:", index)
        assert index >= 1
        if index == 1: # top of the heap
            return
        parent_index = index // 2
        if self.H[parent_index] < self.H[index]:
            return
        else:
            self.H[parent_index], self.H[index] = self.H[index], self.H[parent_index]
            self.bubble_up(parent_index)

    ## bubble_down function at index
    ## WARNING: this function has been cut and paste for the next problem as well
    def bubble_down(self, index):
        assert index >= 1 and index < len(self.H)
        lchild_index = 2 * index
        rchild_index = 2 * index + 1
        # set up the value of left child to the element at that index if valid, or else make it +Infinity
        lchild_value = self.H[lchild_index] if lchild_index < len(self.H) else float('inf')
        # set up the value of right child to the element at that index if valid, or else make it +Infinity
        rchild_value = self.H[rchild_index] if rchild_index < len(self.H) else float('inf')
        # If the value at the index is lessthan or equal to the minimum of two children, then nothing else to do
        if self.H[index] <= min(lchild_value, rchild_value):
            return
        # Otherwise, find the index and value of the smaller of the two children.
        # A useful python trick is to compare
        min_child_value, min_child_index = min ((lchild_value, lchild_index), (rchild_value, rchild_index))
        # Swap the current index with the least of its two children
        self.H[index], self.H[min_child_index] = self.H[min_child_index], self.H[index]
        # Bubble down on the minimum child index
        self.bubble_down(min_child_index)

    # Function: heap_insert
    # Insert elt into heap
    # Use bubble_up/bubble_down function
    def insert(self, elt):
        # print("Hello from heap insert!")
        # your code here
        sh_length = len(self.H)
        # print(f"self.size: {self.size()}")
        # print(f"insert: self.H={self.H}, length={sh_length} elt={elt}")
        self.H.append(elt)
        self.bubble_up(sh_length)
        #self.heapify()


    # Function: heap_delete_min
    # delete the smallest element in the heap. Use bubble_up/bubble_down
    def delete_min(self):  # I really need to figure out how to refactor this
        # your code here
        # print("self.H before: ", self.H)
        # print(f"self.min_element(): {self.min_element()}")

        # print(f"self.size() = {self.size()}")
        if self.size() == 1:
            self.H.pop(1)
        if self.size() > 1:
          item = self.H[1]
          self.H[1] = self.H.pop((self.size()))
          self.bubble_down(1)
          #print("self.H after: ", self.H)
          return item

```

<!-- grade_id: cell-6956a1e5f9f6578e, points: 10, grade: true, solution: false -->
```python
h = MinHeap()
assert(h.is_empty())
print('Inserting: 5, 2, 4, -1 and 7 in that order.')
h.insert(5)
print(f'\t Heap = {h}')
assert(h.min_element() == 5)
h.insert(2)
print(f'\t Heap = {h}')
assert(h.min_element() == 2)
h.insert(4)
print(f'\t Heap = {h}')
assert(h.min_element() == 2)
h.insert(-1)
print(f'\t Heap = {h}')
assert(h.min_element() == -1)
h.insert(7)
print(f'\t Heap = {h}')
assert(h.min_element() == -1)
h.satisfies_assertions()

print('Deleting minimum element')
h.delete_min()
print(f'\t Heap = {h}')
assert(h.min_element() == 2)
h.delete_min()
print(f'\t Heap = {h}')
assert(h.min_element() == 4)
h.delete_min()
print(f'\t Heap = {h}')
assert(h.min_element() == 5)
h.delete_min()
print(f'\t Heap = {h}')
assert(h.min_element() == 7)
# Test delete_max? on heap of size 1, should result in empty heap.
h.delete_min()
assert(h.is_empty())
print(f'\t Heap = {h}')
print('All tests passed: 10 points!')

```

Output:
```
IS EMPTY? [None]
Inserting: 5, 2, 4, -1 and 7 in that order.
	 Heap = [None, 5]
	 Heap = [None, 2, 5]
	 Heap = [None, 2, 5, 4]
	 Heap = [None, -1, 2, 4, 5]
	 Heap = [None, -1, 2, 4, 5, 7]
Deleting minimum element
self.size() = 5
	 Heap = [None, 2, 5, 4, 7]
self.size() = 4
	 Heap = [None, 4, 5, 7]
self.size() = 3
	 Heap = [None, 5, 7]
self.size() = 2
	 Heap = [None, 7]
self.size() = 1
IS EMPTY? [None]
	 Heap = [None]
All tests passed: 10 points!
```

<!-- grade_id: cell-412c6a13c712d06b, points: 0, grade: false, solution: true -->
```python
class TopKHeap:

    # The constructor of the class to initialize an empty data structure
    def __init__(self, k):
        self.k = k
        self.A = []
        self.H = MinHeap()

    def size(self):
        return len(self.A) + (self.H.size())

    def get_jth_element(self, j):
        assert 0 <= j < self.k-1
        assert j < self.size()
        return self.A[j]


    def satisfies_assertions(self):
        # is self.A sorted
        for i in range(len(self.A) -1 ):
            assert self.A[i] <= self.A[i+1], f'Array A fails to be sorted at position {i}, {self.A[i], self.A[i+1]}'
        # is self.H a heap (check min-heap property)
        self.H.satisfies_assertions()
        # is every element of self.A less than or equal to each element of self.H
        for i in range(len(self.A)):
            assert self.A[i] <= self.H.min_element(), f'Array element A[{i}] = {self.A[i]} is larger than min heap element {self.H.min_element()}'



    # Function : insert_into_A
    # This is a helper function that inserts an element `elt` into `self.A`.
    # whenever size is < k,
    #       append elt to the end of the array A
    # Move the element that you just added at the very end of
    # array A out into its proper place so that the array A is sorted.
    # return the "displaced last element" jHat (None if no element was displaced)
    def insert_into_A(self, elt):
        # print("Hello from insert into A!")
        # print("k = ", self.k, "size = ", self.size())
        self.A.append(elt)

        if (len(self.A) >=  self.k):
          j = len(self.A)-1
          while (j >= 1 and self.A[j] < self.A[j-1]):
              # Swap A[j] and A[j-1]
              (self.A[j], self.A[j-1]) = (self.A[j-1], self.A[j])
              j = j -1

          if self.A[j] != None and len(self.A) > self.k:
            return self.A.pop() # this is jHat


    # Function: insert -- insert an element into the data structure.
    # Code to handle when self.size < self.k is already provided

    # FROM VIDEO PRIORITY, PRIORITY QUEUES, and HEAPSORT
    # Algo for insert:
      # 1. Insert at end
      # 2. BubbleUp n + 1 (the item at the end can only be a leaf, and
      #     only need to be bubbled up).

    def insert(self, elt):
        size = self.size()
        # If we have fewer than k elements, handle that in a special manner
        if size <= self.k:
            self.insert_into_A(elt)
            return
        # Code up your algorithm here.
        max = (self.A[len(self.A) - 1]) # largest item in A
        # print(f"elt: {elt} | max: {max}")
        if elt < max:
            # print(f'self.A: {self.A, len(self.A)}')

            result = self.insert_into_A(elt)
            # print(f"self.A after insert: {self.A} result: {result}")

            # take the last item and add it to the heap
            self.H.insert(result)
            # print(f"self.H after insert:    {self.H}")

            # Because it goes in last, is BubbleUp always used?
            # always a leaf, so yes, always bubble up
            self.H.bubble_up(self.H.size())
            # print(f"self.H after bubble_up: {self.H}")



    # Function: Delete top k -- delete an element from the array A
    # In particular delete the j^{th} element where j = 0 means the least element.
    # j must be in range 0 to self.k-1

    # WHICH THEN MEANS A is ready for the next lowest item from H

    # BubbleUp if (An) in jth pos is < parent
    # BubbleDown if An > one or both children

    # FROM VIDEO PRIORITY, PRIORITY QUEUES, and HEAPSORT
    # Algo for delete:
    # Assume array A: [A0...Aj...An]
    # 1: Replace Aj with An
    # 2: Adjust size to n - 1
    # 3: Fix what's broken:
          # Bubble up if An is < parent
          # Bubble down if An > child


    def delete_top_k(self, j):
        k = self.k
        # print(f"*** k: {k}, j: {j}, self.size(): {self.size()}, self.H: {self.H}")
        # print(f"self.A[j] = {self.A[j]}")

        assert self.size() > k # we need not handle the case when size is less than or equal to k
        assert j >= 0
        assert j < self.k
        # your code here

        self.A.pop(j)
        # print(f"d_t_k self.A.pop a looks like: {self.A}")
        # print("self.H.delete_min():" ,self.H.delete_min())
        transfer = self.H.delete_min()
        #print("transfer: ", transfer)
        self.insert_into_A(transfer)

        # Insert the minimum heap element at position  𝑘−1  of the array A.
        # Delete the element at index 1 of the heap

        #print (f"post pop self.A: {self.A}, self.H: {self.H}")
```

<!-- grade_id: cell-2d85778743b67291, points: 15, grade: true, solution: false -->
```python
h = TopKHeap(5) # This passes in 5 as k
# Force the array A
h.A = [-10, -9, -8, -4, 0]
# Force the heap to this heap # What does this mean?
[h.H.insert(elt) for elt in  [1, 4, 5, 6, 15, 22, 31, 7]]

print('Initial data structure: ')
print('\t A = ', h.A)
print('\t H = ', h.H)

# Insert an element -2
print('Test 1: Inserting element -2')
h.insert(-2)
print('\t A = ', h.A)
print('\t H = ', h.H)
# After insertion h.A should be [-10, -9, -8, -4, -2]
# After insertion h.H should be [None, 0, 1, 5, 4, 15, 22, 31, 7, 6]
assert h.A == [-10,-9,-8,-4,-2]
assert h.H.min_element() == 0 , 'Minimum element of the heap is no longer 0'
h.satisfies_assertions()

print('Test2: Inserting element -11')
h.insert(-11)
print('\t A = ', h.A)
print('\t H = ', h.H)
assert h.A == [-11, -10, -9, -8, -4]
assert h.H.min_element() == -2
h.satisfies_assertions()
# %debug
print('Test 3 delete_top_k(3)')
h.delete_top_k(3)
print('\t A = ', h.A)
print('\t H = ', h.H)
h.satisfies_assertions()
assert h.A == [-11,-10,-9,-4,-2]
assert h.H.min_element() == 0
h.satisfies_assertions()

print('Test 4 delete_top_k(4)')
h.delete_top_k(4)
print('\t A = ', h.A)
print('\t H = ', h.H)
assert h.A == [-11, -10, -9, -4, 0]
h.satisfies_assertions()

print('Test 5 delete_top_k(0)')
h.delete_top_k(0)
print('\t A = ', h.A)
print('\t H = ', h.H)
assert h.A == [-10, -9, -4, 0, 1]
h.satisfies_assertions()

print('Test 6 delete_top_k(1)')
h.delete_top_k(1)
print('\t A = ', h.A)
print('\t H = ', h.H)
assert h.A == [-10, -4, 0, 1, 4]
h.satisfies_assertions()
print('All tests passed - 15 points!')

```

Output:
```
Initial data structure: 
	 A =  [-10, -9, -8, -4, 0]
	 H =  [None, 1, 4, 5, 6, 15, 22, 31, 7]
Test 1: Inserting element -2
	 A =  [-10, -9, -8, -4, -2]
	 H =  [None, 0, 1, 5, 4, 15, 22, 31, 7, 6]
Test2: Inserting element -11
	 A =  [-11, -10, -9, -8, -4]
	 H =  [None, -2, 0, 5, 4, 1, 22, 31, 7, 6, 15]
Test 3 delete_top_k(3)
	 A =  [-11, -10, -9, -4, -2]
	 H =  [None, 0, 1, 5, 4, 15, 22, 31, 7, 6]
Test 4 delete_top_k(4)
	 A =  [-11, -10, -9, -4, 0]
	 H =  [None, 1, 4, 5, 6, 15, 22, 31, 7]
Test 5 delete_top_k(0)
	 A =  [-10, -9, -4, 0, 1]
	 H =  [None, 4, 6, 5, 7, 15, 22, 31]
Test 6 delete_top_k(1)
	 A =  [-10, -4, 0, 1, 4]
	 H =  [None, 5, 6, 22, 7, 15, 31]
All tests passed - 15 points!
```

- (a) a sorted array A of  𝑘  elements that forms our least k elements; and **A = least k**
- (b) a minheap H with the remaining  𝑛−𝑘  elements. **H = remainder**

* First, in order to delete the index j from array, move elements from j+1 .. k-1 left one position.
* Insert the minimum heap element at position  𝑘−1  of the array A.
* Delete the element at index 1 of the heap.



## Problem 2: Heap data structure to mantain/extract median (instead of minimum/maximum key)

We have seen how min-heaps can efficiently extract the smallest element efficiently and maintain the least element as we insert/delete elements. Similarly, max-heaps can maintain the largest element. In this exercise, we combine both to maintain the "median" element.

The median is the middle element of a list of numbers.
- If the list has size $n$ where $n$ is odd, the median is the $(n-1)/2^{th}$ element where $0^{th}$ is least and $(n-1)^{th}$ is the maximum.
- If $n$ is even, then we designate the median the average of the $(n/2-1)^{th}$ and $(n/2)^{th}$ elements.


#### Example
Are both of these examples wrong?  
~~- List is $[-1, 5, 4, 2, 3]$ has size $5$, the median is the $2^{nd}$ element (remember again least element is designated as $0^{th}$) which is $3$.~~
~~- List is $[-1, 3, 2, 1 ]$ has size $4$. The median element is the average of  $1^{st}$ element (1) and $2^{nd}$ element (2) which is  $1.5$.~~
- List $[-1, 5, 4, 2, 3]$ has size $5$, the median is the $2^{nd}$ element (counting from $0$) which is $4$.  
- List $[-1, 3, 2, 1]$ has size $5$; the median is the average of the $1^{st}$ and $2^{nd}$ elements (again counting from $0$) which is (3 + 2) / 2 = 2.5.

## Maintaining median using two heaps.

The data will be maintained as the union of the elements in two heaps $H_{\min}$ and $H_{\max}$, wherein $H_{\min}$ is a min-heap and $H_{\max}$ is a max-heap.  We will maintain the following invariant:
  - The max element of  $H_{\max}$ will be less than or equal to the min element of  $H_{\min}$.
  - The sizes of $H_{max}$ and $H_{min}$ are equal (if number of elements in the data structure is even) or $H_{max}$ may have one less element than $H_{min}$ (if the number of elements in the data structure is odd).
  


## (A)  Design algorithm for insertion.

Suppose, we have the current data split between $H_{max}$ and $H_{min}$ and we wish to insert an element $e$ into the data structure, describe the algorithm you will use to insert. Your algorithm must decide which of the two heaps will $e$ be inserted into and how to maintain the size balance condition.

Describe the algorithm below and the overall complexity of an insert operation. This part will not be graded.

<!-- grade_id: cell-3366a32bb62b3cc3, points: 0, grade: true, solution: true -->
YOUR ANSWER HERE

1.   Figure out which list it belongs in
2.   Put it there
3.   If needed, insert / remove to maintain the balance



## (B) Design algorithm for finding the median.

Implement an algorithm for finding the median given the heaps $H_{\min}$ and $H_{\max}$. What is its complexity?

<!-- grade_id: cell-99c28c7d7bb49c98, points: 0, grade: true, solution: true -->
YOUR ANSWER HERE

## (C) Implement the algorithm

Complete the implementation for maxheap data structure.
First complete the implementation of MaxHeap.  You can cut and paste relevant parts from previous problems although we do not really recommend doing that. A better solution is to write a single implementation that can serve as either a  min or max heap based on a flag.

```python
# Start the iPdb Debugger
# !pip install -Uqq ipdb
# import ipdb
# %pdb on
%pdb off
```

Output:
```
Automatic pdb calling has been turned OFF
```

## `class MinHeap`

```python
# First let us complete a minheap data structure.
# Please complete missing parts below.

class MinHeap:
    def __init__(self):
        # H = the heap of the N-K elements.
        self.H = [None]

    def size(self):
        return len(self.H)-1

    def __repr__(self):
        return str(self.H[0:])

    def satisfies_assertions(self):
        for i in range(2, len(self.H)):
            assert self.H[i] >= self.H[i//2],  f'Min heap property fails at position {i//2}, parent elt: {self.H[i//2]}, child elt: {self.H[i]}'

    def min_element(self):
        return self.H[1]

    def bubble_up(self, index):
        assert index >= 1
        if index == 1: # top of the heap
            return
        parent_index = index // 2
        if self.H[parent_index] < self.H[index]:
            return
        else:
            self.H[parent_index], self.H[index] = self.H[index], self.H[parent_index]
            self.bubble_up(parent_index)

    def bubble_down(self, index):
        assert index >= 1 and index < len(self.H)
        lchild_index = 2 * index
        rchild_index = 2 * index + 1
        # set up the value of left child to the element at that index if valid, or else make it +Infinity
        lchild_value = self.H[lchild_index] if lchild_index < len(self.H) else float('inf')
        # set up the value of right child to the element at that index if valid, or else make it +Infinity
        rchild_value = self.H[rchild_index] if rchild_index < len(self.H) else float('inf')
        # If the value at the index is lessthan or equal to the minimum of two children, then nothing else to do
        if self.H[index] <= min(lchild_value, rchild_value):
            return
        # Otherwise, find the index and value of the smaller of the two children.
        # A useful python trick is to compare
        min_child_value, min_child_index = min ((lchild_value, lchild_index), (rchild_value, rchild_index))
        # Swap the current index with the least of its two children
        self.H[index], self.H[min_child_index] = self.H[min_child_index], self.H[index]
        # Bubble down on the minimum child index
        self.bubble_down(min_child_index)

    # Function: heap_insert
    # Insert elt into heap
    # Use bubble_up/bubble_down function
    def insert(self, elt):
        # print("Hello from heap insert!")
        # your code here
        sh_length = len(self.H)
        # print(f"self.size: {self.size()}")
        # print(f"insert: self.H={self.H}, length={sh_length} elt={elt}")
        self.H.append(elt)
        self.bubble_up(sh_length)
        #self.heapify()


    # Function: heap_delete_min
    # delete the smallest element in the heap. Use bubble_up/bubble_down
    def delete_min(self):
        # your code here
        if self.size() == 1:
            self.H.pop(1)
        if self.size() > 1:
          item = self.H[1]
          self.H[1] = self.H.pop((self.size()))
          self.bubble_down(1)
          return item

```

## `class MaxHeap`

<!-- grade_id: cell-4d474251918c6d96, points: 0, grade: false, solution: true -->
```python

class MaxHeap:
    def __init__(self):
        self.H = [None] # don't forget to maintain H[0] as None

    def size(self):
        return len(self.H)-1

    def __repr__(self):
        return str(self.H[0:])

    def satisfies_assertions(self):
        for i in range(2, len(self.H)):
            assert self.H[i] <= self.H[i//2],  f'Maxheap property fails at position {i//2}, parent elt: {self.H[i//2]}, child elt: {self.H[i]}'

    def max_element(self):
        return self.H[1]

    def bubble_up(self, index):
        # print("HELLO FROM BUBBLE_UP at index:", index)
        assert index >= 1

        if index == 1: # top of the heap
            return
        parent_index = index // 2


        if self.H[parent_index] > self.H[index]:
            return
        else:
            self.H[parent_index], self.H[index] = self.H[index], self.H[parent_index]
            self.bubble_up(parent_index)



    def bubble_down(self, index):
        assert index >= 1 and index < len(self.H)
        lchild_index = 2 * index
        rchild_index = 2 * index + 1
        # set up the value of left child to the element at that index if valid, or else make it -Infinity
        lchild_value = self.H[lchild_index] if lchild_index < len(self.H) else float('-inf')
        # set up the value of right child to the element at that index if valid, or else make it -Infinity
        rchild_value = self.H[rchild_index] if rchild_index < len(self.H) else float('-inf')
        # If the value at the index is greater than or equal to the maximum of two children, then nothing else to do
        if self.H[index] >= max(lchild_value, rchild_value):
            return
        # Otherwise, find the index and value of the larger of the two children.
        # A useful python trick is to compare
        max_child_value, max_child_index = max((lchild_value, lchild_index), (rchild_value, rchild_index))
        # Swap the current index with the greatest of its two children
        self.H[index], self.H[max_child_index] = self.H[max_child_index], self.H[index]
        # Bubble down on the maximum child index
        self.bubble_down(max_child_index)


    # Function: insert
    # Insert elt into minheap
    # Use bubble_up/bubble_down function
    def insert(self, elt):
        # your code here
        # self.bubble_up(10, True)
        # self.bubble_up(10, False)
        self.H.append(elt)
        self.bubble_up(self.size())

    # Function: delete_max
    # delete the largest element in the heap. Use bubble_up/bubble_down
    # def delete_max(self):
    #     # your code here
    #     print("HELLO FROM DELETE_MAX")

    # delete the largest element in the heap. Use bubble_up/bubble_down
    def delete_max(self): # This seems like it could use a refactor, but I'm not sure how...
        # your code here
        if self.size() == 1:
          self.H.pop(1)
        if self.size() > 1:
          item = self.H[1]
          self.H[1] = self.H.pop((self.size()))
          self.bubble_down(1)
          return item

```

### `MaxHeap` test

<!-- grade_id: cell-0f15a9f966d07a21, points: 5, grade: true, solution: false -->
```python
h = MaxHeap()
print('Inserting: 5, 2, 4, -1 and 7 in that order.')
h.insert(5)
print(f'\t Heap = {h}')
assert(h.max_element() == 5)
h.insert(2)
print(f'\t Heap = {h}')
assert(h.max_element() == 5)
h.insert(4)
print(f'\t Heap = {h}')
assert(h.max_element() == 5)
h.insert(-1)
print(f'\t Heap = {h}')
assert(h.max_element() == 5)
h.insert(7)
print(f'\t Heap = {h}')
assert(h.max_element() == 7)
h.satisfies_assertions()

print('Deleting maximum element')
h.delete_max()
print(f'\t Heap = {h}')
assert(h.max_element() == 5)
h.delete_max()
print(f'\t Heap = {h}')
assert(h.max_element() == 4)
h.delete_max()
print(f'\t Heap = {h}')
assert(h.max_element() == 2)
h.delete_max()
print(f'\t Heap = {h}')
assert(h.max_element() == -1)
# Test delete_max on heap of size 1, should result in empty heap.
h.delete_max()
print(f'\t Heap = {h}')
print('All tests passed: 5 points!')

```

Output:
```
Inserting: 5, 2, 4, -1 and 7 in that order.
	 Heap = [None, 5]
	 Heap = [None, 5, 2]
	 Heap = [None, 5, 2, 4]
	 Heap = [None, 5, 2, 4, -1]
	 Heap = [None, 7, 5, 4, -1, 2]
Deleting maximum element
	 Heap = [None, 5, 2, 4, -1]
	 Heap = [None, 4, 2, -1]
	 Heap = [None, 2, -1]
	 Heap = [None, -1]
	 Heap = [None]
All tests passed: 5 points!
```

## class MedianMaintainingHeap

<!-- grade_id: cell-91adccbf6dbb6a77, points: 0, grade: false, solution: true -->
```python
class MedianMaintainingHeap:
    def __init__(self):
        self.hmin = MinHeap()
        self.hmax = MaxHeap()

    def satisfies_assertions(self):
        # print("SATISFIES ASSERTIONS?")
        if self.hmin.size() == 0:
            assert self.hmax.size() == 0
            return
        if self.hmax.size() == 0:
            assert self.hmin.size() == 1
            return
        # 1. min heap min element >= max heap max element
        assert self.hmax.max_element() <= self.hmin.min_element(), f'Failed: Max element of max heap = {self.hmax.max_element()} > min element of min heap {self.hmin.min_element()}'
        # 2. size of max heap must be equal or one lessthan min heap.
        s_min = self.hmin.size()
        s_max = self.hmax.size()
        assert (s_min == s_max or s_max  == s_min -1 ), f'Heap sizes are unbalanced. Min heap size = {s_min} and Maxheap size = {s_max}'

    def __repr__(self):
        return 'Maxheap:' + str(self.hmax) + ' Minheap:'+str(self.hmin)

    def get_median(self):
        # print("GET_MEDIAN")
        if self.hmin.size() == 0:
            assert self.hmax.size() == 0, 'Sizes are not balanced'
            assert False, 'Cannot ask for median from empty heaps'
        if self.hmax.size() == 0:
            assert self.hmin.size() == 1, 'Sizes are not balanced'
            return self.hmin.min_element()
        # your code here
        #  If sizes of heaps are the same:
        #  median = average of self.hmax.max_element +  self.hmin.min_element / 2
        if self.hmax.size() == self.hmin.size():
          median = (self.hmax.max_element() + self.hmin.min_element()) / 2
          # print(f"SAME SIZE HEAPS MEDIAN: {median}")
          return median
        #  otherwise:
        #  median = self.hmin.min_element
        else: # for what other case could there be?
          median = self.hmin.min_element()
          # print(f"BIGGER SIZE HMIN HEAP ELEMENT {median}")
          return median


    # function: balance_heap_sizes
    # ensure that the size of hmax == size of hmin or size of hmax +1 == size of hmin
    # If the condition above does not hold, move the max element from max heap into the min heap or
    # vice versa as needed to balance the sizes.
    # This function could be called from insert/delete_median methods
    def balance_heap_sizes(self):
        print("BALANCE HEAP SIZES")
        # your code here
        print(f"self.hmin.size(): {self.hmin.size()} | self.hmax.size(): {self.hmax.size()}")
        # sample max_heap=[7, 5, 4, 2, -1] sample min_heap=[10, 15, 20, 30, 45]
        # sample unbalanced heaps: Maxheap:[1] Minheap:[2, 5, 4]
        # If  𝐻max  is larger then, extract the largest element from  𝐻max  andd insert into  𝐻min .
        if self.hmax.size() > self.hmin.size():
          self.hmin.insert(self.hmax.delete_max())
        # If  𝐻min  is larger then, extract the least element from  𝐻min  and insert into  𝐻max .
        else:
          self.hmax.insert(self.hmin.delete_min())

    def insert(self, elt):
        # Handle the case when either heap is empty
        if self.hmin.size() == 0:
            # min heap is empty -- directly insert into min heap
            print(f"inserting {elt} into h_min")
            self.hmin.insert(elt)
            return
        if self.hmax.size() == 0:
            # max heap is empty -- this better happen only if min heap has size 1.
            assert self.hmin.size() == 1
            if elt > self.hmin.min_element():
                # Element needs to go into the min heap
                current_min = self.hmin.min_element()
                print(f"removing {current_min} from h_min and inserting it into h_max")
                self.hmin.delete_min()
                self.hmax.insert(current_min)
                print(f"...and inserting {elt} into h_min")
                self.hmin.insert(elt)
                # done!
            else:
                # Element goes into the max heap -- just insert it there.
                print(f"inserting {elt} into h_max")
                self.hmax.insert(elt)
            return
        # Now assume both heaps are non-empty
        # your code here
        # Let  𝑎  be the largest element in  𝐻max  and  𝑏  be the least element in  𝐻min .
        a = self.hmax.max_element()
        b = self.hmin.min_element()
        #If  𝑒𝑙𝑡<𝑎 , then we insert the new element into  𝐻max .
        if elt < a:
          self.hmax.insert(elt)

        #If  𝑒𝑙𝑡>=𝑎 , then we insert the new element into  𝐻min .
        if elt >= a:
          self.hmin.insert(elt)

        # If the size of  𝐻max  and  𝐻min  differ by 2, then
        if abs(self.hmax.size() - self.hmin.size()) >= 2:
          self.balance_heap_sizes()




    def delete_median(self):
        self.hmax.delete_max()
        self.balance_heap_sizes()
```

## `MedianMaintainingHeap()`

<!-- grade_id: cell-dbf8d2bf71617a5e, points: 15, grade: true, solution: false -->
```python
m = MedianMaintainingHeap()
print('Inserting 1, 5, 2, 4, 18, -4, 7, 9')

m.insert(1)
print(m)
print(m.get_median())
m.satisfies_assertions()
assert m.get_median() == 1,  f'expected median = 1, your code returned {m.get_median()}'

m.insert(5)
print(m)
print(m.get_median())
m.satisfies_assertions()
assert m.get_median() == 3,  f'expected median = 3.0, your code returned {m.get_median()}'

m.insert(2)
print(m)
print(m.get_median())
m.satisfies_assertions()

assert m.get_median() == 2,  f'expected median = 2, your code returned {m.get_median()}'
m.insert(4)
print(m)
print(m.get_median())
m.satisfies_assertions()
assert m.get_median() == 3,  f'expected median = 3, your code returned {m.get_median()}'

m.insert(18)
print(m)
print(m.get_median())
m.satisfies_assertions()
assert m.get_median() == 4,  f'expected median = 4, your code returned {m.get_median()}'

m.insert(-4)
print(m)
print(m.get_median())
m.satisfies_assertions()
assert m.get_median() == 3,  f'expected median = 3, your code returned {m.get_median()}'

m.insert(7)
print(m)
print(m.get_median())
m.satisfies_assertions()
assert m.get_median() == 4, f'expected median = 4, your code returned {m.get_median()}'

m.insert(9)
print(m)
print(m.get_median())
m.satisfies_assertions()
assert m.get_median()== 4.5, f'expected median = 4.5, your code returned {m.get_median()}'

print('All tests passed: 15 points')

```

Output:
```
Inserting 1, 5, 2, 4, 18, -4, 7, 9
inserting 1 into h_min
Maxheap:[None] Minheap:[None, 1]
1
removing 1 from h_min and inserting it into h_max
...and inserting 5 into h_min
Maxheap:[None, 1] Minheap:[None, 5]
3.0
Maxheap:[None, 1] Minheap:[None, 2, 5]
2
BALANCE HEAP SIZES
self.hmin.size(): 3 | self.hmax.size(): 1
Maxheap:[None, 2, 1] Minheap:[None, 4, 5]
3.0
Maxheap:[None, 2, 1] Minheap:[None, 4, 5, 18]
4
Maxheap:[None, 2, 1, -4] Minheap:[None, 4, 5, 18]
3.0
Maxheap:[None, 2, 1, -4] Minheap:[None, 4, 5, 18, 7]
4
BALANCE HEAP SIZES
self.hmin.size(): 5 | self.hmax.size(): 3
Maxheap:[None, 4, 2, -4, 1] Minheap:[None, 5, 7, 18, 9]
4.5
All tests passed: 15 points
```

## Solutions to Manually Graded Portions

### Problem 1 A
In order to insert a new element `j`, we will first  distinguish between two cases:
- $j < A[k-1]$ : In this case $j$ belongs to the array $A$.
  - First, let $j' = A[k-1]$.
  - Replace $A[k-1]$ by $j$.
  - Perform an insertion to move $j$ into its correct place in the sorted array $A$.
  - Insert $j'$ into the heap using heap insert.
- $j \geq A[k-1]$: In this case, $j$ belongs to the heap $H$.
  - Insert $j$ into the heap using heap-insert.
  
In terms of $k, n$, the worst case complexity is $\Theta(k + \log(n))$ for each insertion operation.

### Problem 1B

- First, in order to delete the index j from array, move elements from j+1 .. k-1 left one position.
- Insert the minimum heap element at position $k-1$ of the array A.
- Delete the element at index 1 of the heap.

Overall complexity = $\Theta(k + \log(n))$ in the worst case.

### Problem 2 A

sample max_heap=[7, 5, 4, 2, -1] sample min_heap=[10, 15, 20, 30, 45]

Let $a$ be the largest element in $H_{\max}$ and $b$ be the least element in $H_{\min}$.
 - If $elt < a$, then we insert the new element into $H_{\max}$.
 - If $elt >= a$, then we insert the new element into $H_{\min}$.

 If the size of $H_{\max}$ and $H_{\min}$ differ by 2, then
 - If $H_{\max}$ is larger then, extract the largest element from $H_{\max}$ andd insert into $H_{\min}$.
 - If $H_{\min}$ is larger then,  extract the least element from $H_{\min}$ andd insert into $H_{\max}$.

 The overall complexity is $\Theta(\log(n))$.

### Problem 2 B

 If sizes of heaps are the same, then median is the average of maximum element of the max heap and minimum element of the minheap.

Otherwise, the median is simply the minimum elemment of the min-heap.

Overall complexity is $\Theta(1)$.



## That's all folks